```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

# Lab 6 – APIs and Web Scraping

## DSC 80, Fall 2025

### Due Date: Monday, November 10th at 11:59 PM


## Instructions

Welcome to the sixth DSC 80 lab this quarter!

Much like in DSC 10, this Jupyter Notebook contains the statements of the problems and provides code and Markdown cells to display your answers to the problems. Unlike DSC 10, the notebook is *only* for displaying a readable version of your final answers. The coding will be done in an accompanying `lab.py` file that is imported into the current notebook, and **you will only submit that `lab.py` file**, not this notebook!

Some additional guidelines:
- **Unlike in DSC 10, labs will have both public tests and hidden tests.** The bulk of your grade will come from your scores on hidden tests, which you will only see on Gradescope after the assignment deadline.
- **Do not change the function names in the `lab.py` file!** The functions in the `lab.py` file are how your assignment is graded, and they are graded by their name. If you changed something you weren't supposed to, you can find the original code in the [course GitHub repository](https://github.com/dsc-courses/dsc80-2025-sp).
- Notebooks are nice for testing and experimenting with different implementations before designing your function in your `lab.py` file. You can write code here, but make sure that all of your real work is in the `lab.py` file, since that's all you're submitting.
- You are encouraged to write your own additional helper functions to solve the lab, as long as they also end up in `lab.py`.

**To ensure that all of the work you want to submit is in `lab.py`, we've included a script named `lab-validation.py` in the lab folder. You shouldn't edit it, but instead, you should call it from the command line (e.g. the Terminal) to test your work.** More details on its usage are given at the bottom of this notebook.

**Importing code from `lab.py`**:

* Below, we import the `.py` file that's contained in the same directory as this notebook.
* We use the `autoreload` notebook extension to make changes to our `lab.py` file immediately available in our notebook. Without this extension, we would need to restart the notebook kernel to see any changes to `lab.py` in the notebook.
    - `autoreload` is necessary because, upon import, `lab.py` is compiled to bytecode (in the directory `__pycache__`). Subsequent imports of `lab` merely import the existing compiled python.

<div class="alert alert-block alert-danger">
<b>Note: </b> For this lab, due to system constraints, the code may take varying amounts of time to run on different machines. If you are able to submit to Gradescope and the code runs, then you have met the runtime requirements.
</div>

```python
%load_ext autoreload
%autoreload 2
```

Output:
```
The autoreload extension is already loaded. To reload it, use:
  %reload_ext autoreload
```

```python
from lab import *
```

If the cell below returns a `ModuleNotFoundError`, please run `!pip install lxml` in a new cell. After `lxml` is succesfully installed, go to `kernel` then restart. Note that you will only need to run `!pip install lxml` once. 

```python
import os
import pandas as pd
import numpy as np
import requests
import bs4
import lxml
```

```python
#!pip install lxml
```

## Question 1 – Practice with HTML Tags 📎

In Question 2, you'll spend plenty of time parsing HTML source code. But before you get your hands dirty trying to extract information from HTML written by other people, it is a good idea to write basic HTML code yourself. This exercise will help you better understand how the code in a `.html` file is structured.

For this question, you'll create a very basic `.html` file, named `lab06_1.html`, that satisfies the following conditions:

- It must have `<title>` and `<head>` tags.
- It must also have `<body>` tags. Within the `<body>` tags, it must have:
    - At least two headers.
    * At least three images.
        - At least one image must be a local file.
        - At least one image must be linked to online source.
        - At least one image has to have default text when it cannot be displayed.
    * At least three references (hyperlinks) to different web pages.
    * At least one table with two rows and two columns.
    

Make sure to save your file as `lab06_1.html`, and save it in the same directory as `lab.py`. **When submitting this homework to Gradescope, make sure to also upload `lab06_1.html` along with the local image that you embedded in your site.** You can upload multiple files to Gradescope at a time.
   

***Notes***:
- You can write and view basic HTML with a Jupyter Notebook, using either a Markdown cell or by using the `IPython.display.HTML` function (which takes in a string of HTML and renders it).
- If you write your HTML code within a Jupyter Notebook, you should later copy your code into a text editor and save it with the `.html` extension. You could also write your HTML in a text editor directly.
- Be sure to open your final `.html` file in a browser and make sure it looks correct on its own.

```python

```

```python

```

```python
# Don't delete this cell!
question1()
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

## Question 2 – Scraping an Online Bookstore 📚

Browse through the following fake online bookstore: http://books.toscrape.com/. This website is meant for toying with scraping.

Your job is to scrape the website, collecting data on all books that have:
- **_at least_ a four-star rating**, and
- **a price _strictly_ less than £50**, and 
- **belong to specific categories** (more details below). 

You will extract the information into a DataFrame that looks like the one below.

<table border="1" class="dataframe">
  <thead>
    <tr style="text-align: right;">
      <th></th>
      <th>UPC</th>
      <th>Product Type</th>
      <th>Price (excl. tax)</th>
      <th>Price (incl. tax)</th>
      <th>Tax</th>
      <th>Availability</th>
      <th>Number of reviews</th>
      <th>Category</th>
      <th>Rating</th>
      <th>Description</th>
      <th>Title</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <th>0</th>
      <td>e10e1e165dc8be4a</td>
      <td>Books</td>
      <td>Â£22.60</td>
      <td>Â£22.60</td>
      <td>Â£0.00</td>
      <td>In stock (19 available)</td>
      <td>0</td>
      <td>Default</td>
      <td>Four</td>
      <td>For readers of Laura Hillenbrand's Seabiscuit...</td>
      <td>The Boys in the Boat: Nine Americans...</td>
    </tr>
    <tr>
      <th>1</th>
      <td>c2e46a2ee3b4a322</td>
      <td>Books</td>
      <td>Â£25.27</td>
      <td>Â£25.27</td>
      <td>Â£0.00</td>
      <td>In stock (19 available)</td>
      <td>0</td>
      <td>Romance</td>
      <td>Five</td>
      <td>A Michelin two-star chef at twenty-eight, Violette...</td>
      <td>Chase Me (Paris Nights #2)</td>
    </tr>
    <tr>
      <th>2</th>
      <td>00bfed9e18bb36f3</td>
      <td>Books</td>
      <td>Â£34.53</td>
      <td>Â£34.53</td>
      <td>Â£0.00</td>
      <td>In stock (19 available)</td>
      <td>0</td>
      <td>Romance</td>
      <td>Five</td>
      <td>No matter how busy he keeps himself...</td>
      <td>Black Dust</td>
    </tr>
  </tbody>
</table>

To do so, implement the following functions.

<br>

#### `extract_book_links`

Complete the implementation of the function `extract_book_links`, which takes in the content of a page that contains book listings as a **string of HTML**, and returns a **list** of URLs of book-specific pages for all books with **_at least_ a four-star rating and a price _strictly_ less than £50**.

For this method, the URLs you return should not contain the protocol (i.e. `'https://'`). The protocols should be added back into the URLs when you actually make the requests.


<br>

#### `get_product_info`

Complete the implementation of the function `get_product_info`, which takes in the content of a book-specific page as a **string of HTML**, and a list `categories` of book categories. If the input book is in the list of `categories`, `get_product_info` should return a dictionary corresponding to a row in the DataFrame in the image above (where the keys are the column names and the values are the row values). If the input book is not in the list of `categories`, return `None`. <b> The order of the dictionary keys does not matter. </b>


<br>

#### `scrape_books`

Finally, put everything together. Complete the implementation of the function `scrape_books`, which takes in an integer `k` and a list `categories` of book categories. `scrape_books` should use `requests` to scrape the first `k` pages of the bookstore and return a DataFrame of only the books that have:
- **_at least_ a four-star rating**, and
- **a price _strictly_ less than £50**, and
- **a category that is in the list `categories`**.

<b> The order of the DataFrame columns does not matter. </b>

<br>

Some general guidance and tips:

- The first page of the bookstore is at http://books.toscrape.com/catalogue/page-1.html. Subsequent pages can be found by clicking the "Next" button at the bottom of the page. Look at how the URLs change each time you navigate to a new page; think about how to use f-strings (or some other string formatting technique) to generate these URLs.
- Use "inspect element" to view the source code of the pages you're trying to scrape. To find a book's category, look at the hyperlinks in the book-specific page for that book.
- **`scrape_books` should run in under 180 seconds on the entire bookstore (`k = 50`). `scrape_books` is also the only function that should make `GET` requests; the other two functions parse already-existing HTML.**
- When instantiating `bs4.BeautifulSoup` objects, use the optional argument `features='lxml'` to suppress any warnings.
- Don't worry about typecasting, i.e. it's fine if `'Number of reviews'` is not stored as type `int`. Also, don't worry if you run into encoding errors in your price columns (as the example DataFrame at the top of this cell contains).

```python
# Part 1 rough work
soup = bs4.BeautifulSoup(open(extract_book_links_fp, encoding='utf-8').read())
res = []
for product_pod in soup.find_all('article', class_='product_pod'):
    price = float(product_pod.find('p', class_='price_color').text[1:])
    rating = product_pod.find('p').get('class')[1]
    if rating in ['Four', 'Five'] and price < 50:
        res.append(product_pod.find('a').get('href'))
res
```

Output:
```
['seven-brief-lessons-on-physics_219/index.html',
 'scarlet-the-lunar-chronicles-2_218/index.html',
 'saga-volume-3-saga-collected-editions-3_216/index.html',
 'running-with-scissors_215/index.html',
 'rise-of-the-rocket-girls-the-women-who-propelled-us-from-missiles-to-the-moon-to-mars_213/index.html',
 'ready-player-one_209/index.html']
```

```python
# Part 2 rough work
soup = bs4.BeautifulSoup(open(get_product_info_fp, encoding='utf-8').read())
category = soup.find('ul', class_='breadcrumb').find_all('a')[-1].text
product_type = soup.find('ul', class_='breadcrumb').find_all('a')[1].text
upc = soup.find('th', string='UPC').find_next_sibling('td').text
price_excl_tax = soup.find('th', string='Price (excl. tax)').find_next_sibling('td').text
price_incl_tax = soup.find('th', string='Price (incl. tax)').find_next_sibling('td').text
tax = soup.find('th', string='Tax').find_next_sibling('td').text
availability = soup.find('th', string='Availability').find_next_sibling('td').text
num_reviews = soup.find('th', string='Number of reviews').find_next_sibling('td').text
rating = soup.find('p', class_='star-rating').get('class')[1]
description = soup.find('div', id='product_description').find_next_sibling('p').text
title = soup.find('div', class_='product_main').find('h1').text
```

```python
# Part 3 rough work
def extract_book_links(text):
    soup = bs4.BeautifulSoup(text)
    res = []
    for product_pod in soup.find_all('article', class_='product_pod'):
        price = float(product_pod.find('p', class_='price_color').text.strip()[2:])
        rating = product_pod.find('p').get('class')[1]
        if rating in ['Four', 'Five'] and price < 50:
            res.append(product_pod.find('a').get('href'))
    return res

def get_product_info(text, categories):
    soup = bs4.BeautifulSoup(text)
    category = soup.find('ul', class_='breadcrumb').find_all('a')[-1].text.strip()
    if category in categories:
        product_type = soup.find('ul', class_='breadcrumb').find_all('a')[1].text.strip()
        upc = soup.find('th', string=lambda x: x.strip() == 'UPC').find_next_sibling('td').text.strip()
        price_excl_tax = soup.find('th', string=lambda x: x.strip() == 'Price (excl. tax)').find_next_sibling('td').text.strip()
        price_incl_tax = soup.find('th', string=lambda x: x.strip() == 'Price (incl. tax)').find_next_sibling('td').text.strip()
        tax = soup.find('th', string=lambda x: x.strip() == 'Tax').find_next_sibling('td').text.strip()
        availability = soup.find('th', string=lambda x: x.strip() == 'Availability').find_next_sibling('td').text.strip()
        num_reviews = soup.find('th', string=lambda x: x.strip() == 'Number of reviews').find_next_sibling('td').text.strip()
        rating = soup.find('p', class_='star-rating').get('class')[1]
        description = soup.find('div', id='product_description').find_next_sibling('p').text.strip()
        title = soup.find('div', class_='product_main').find('h1').text.strip()
        return {
            'UPC': upc,
            'Product Type': product_type,
            'Price (excl. tax)': price_excl_tax,
            'Price (incl. tax)': price_incl_tax,
            'Tax': tax,
            'Availability': availability,
            'Number of reviews': num_reviews,
            'Category': category,
            'Rating': rating,
            'Description': description,
            'Title': title
        }
    return None

def download_page_num(i):
    url = f'http://books.toscrape.com/catalogue/page-{i}.html'
    request = requests.get(url)
    return bs4.BeautifulSoup(request.text)

def download_page_book(suffix):
    url = f'http://books.toscrape.com/catalogue/{suffix}'
    request = requests.get(url)
    return bs4.BeautifulSoup(request.text)

links = extract_book_links(download_page_num(1).prettify())
for link in links:
    print(get_product_info(download_page_book(link).prettify(), ['Mystery']))
# bs4.BeautifulSoup(download_page_book(links[0]).prettify()).find_all('th', string=lambda x: x.strip() == 'UPC')
```

Output:
```
{'UPC': 'e00eb4fd7b871a48', 'Product Type': 'Books', 'Price (excl. tax)': 'Â£47.82', 'Price (incl. tax)': 'Â£47.82', 'Tax': 'Â£0.00', 'Availability': 'In stock (20 available)', 'Number of reviews': '0', 'Category': 'Mystery', 'Rating': 'Four', 'Description': 'WICKED above her hipbone, GIRL across her heart Words are like a road map to reporter Camille Preakerâ\x80\x99s troubled past. Fresh from a brief stay at a psych hospital, Camilleâ\x80\x99s first assignment from the second-rate daily paper where she works brings her reluctantly back to her hometown to cover the murders of two preteen girls. NASTY on her kneecap, BABYDOLL on her leg Since WICKED above her hipbone, GIRL across her heart Words are like a road map to reporter Camille Preakerâ\x80\x99s troubled past. Fresh from a brief stay at a psych hospital, Camilleâ\x80\x99s first assignment from the second-rate daily paper where she works brings her reluctantly back to her hometown to cover the murders of two preteen girls. NASTY 
```

```python
# don't delete this cell, but do run it -- it is needed for the autograder tests

# public test for extract_book_links 
extract_book_links_fp = os.path.join('data', 'products.html')
extract_book_out = extract_book_links(
    open(extract_book_links_fp, encoding='utf-8').read()
)
extract_book_url = 'scarlet-the-lunar-chronicles-2_218/index.html'

# doc tests for get product info
get_product_info_fp = os.path.join('data', 'Frankenstein.html')
get_product_info_out = get_product_info(
    open(get_product_info_fp, encoding='utf-8').read(), ['Default']
)

# public test for scrape books 
scrape_books_out = scrape_books(1, ['Mystery'])
scrape_books_out
```

Output:
```
c:\Users\sweek\Downloads\ucsd\ucsd_classes\fa25\dsc80\dsc80-2025-fa\labs\lab06\lab.py:81: GuessedAtParserWarning: No parser was explicitly specified, so I'm using the best available HTML parser for this system ("lxml"). This usually isn't a problem, but if you run this code on another system, or in a different virtual environment, it may use a different parser and behave differently.

The code that caused this warning is on line 81 of the file c:\Users\sweek\Downloads\ucsd\ucsd_classes\fa25\dsc80\dsc80-2025-fa\labs\lab06\lab.py. To get rid of this warning, pass the additional argument 'features="lxml"' to the BeautifulSoup constructor.

  return bs4.BeautifulSoup(request.text)
c:\Users\sweek\Downloads\ucsd\ucsd_classes\fa25\dsc80\dsc80-2025-fa\labs\lab06\lab.py:37: GuessedAtParserWarning: No parser was explicitly specified, so I'm using the best available HTML parser for this system ("lxml"). This usually isn't a problem, but if you run this code on another system, or in a different vi
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

## Question 3 – API Requests 🤑

You trade stocks as a hobby. As an avid `pandas` coder, you decide to calculate statistics of your favorite stocks by pulling data from a public API. The API we will work with can be found at https://financialmodelingprep.com/developer/docs/#Stock-Historical-Price. Specifically, we will use the "**Daily Chart EOD**" endpoint (search for it at the linked page).

Some relevant definitions:
- Ticker: A short code that refers to a stock. For example, Apple's ticker is AAPL and Ford's ticker is F. 
- Open: The price of a stock at the beginning of a trading day.
- Close: The price of a stock at the end of a trading day.
- Volume: The total number of shares traded in a day.
- Percent change: The difference in price with respect to the original price, as a percentage.

To make requests to the aforementioned API, you will need an API key. In order to get one, you will need to make an account at the website. Once you've signed up, you can use the API key that comes with the free plan. It has a limit of 250 requests per day, which should be more than enough. You will have to encode your API key in the URL that you make requests to; see a complete example of such a request at the right side of the [documentation](https://site.financialmodelingprep.com/developer/docs#Stock-Historical-Price).

Implement the following two functions.

#### `stock_history`

Complete the implementation of the function `stock_history`, which takes in a string `ticker` and two integers, `year` and `month`, and returns a DataFrame containing the price history for that stock in that month. Keep all of the attributes that are returned by the API.

***Notes***:
- Read the API documentation if you get stuck!
- [`pd.date_range`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.date_range.html) takes in two dates and returns a sequence of all dates between the two dates, excluding the right endpoint. How might this be helpful?
- The [`requests.get`](https://docs.python-requests.org/en/master/user/quickstart/) function returns a Response object, not the data itself. Use the `json` method on the Response object to extract the relevant JSON, as we did in [Lecture 9](https://dsc80.com/resources/lectures/lec09/lec09-filled.html#Example:-GET-requests-via-requests) (you don't need to `import json` to do this).
- You can instantiate a DataFrame using a sequence of dictionaries.

<br>

#### `stock_stats`

Create a function `stock_stats` that takes in a DataFrame outputted by `stock_history` and returns a **tuple** of two numbers:
1. The percent change of the stock throughout the month as a **percentage**.
2. An estimate of the total transaction volume **in billion of dollars** for that month.

Both values in the tuple should be **strings** that contain numbers rounded to two decimal places. Add a plus or minus sign in front of the percent change, and make sure that the total transaction volume string ends in a `'B'`.

**To compute the percent change**, use the opening price on the first day of the month as the starting price and the closing price on the last day of the month as the ending price.

**To compute the total transaction volume**, assume that on any given day, the average price of a share is the midpoint of the high and low price for that day.

$$ \text{Estimated Total Transaction Volume (in dollars)} = \text{Volume (number of shares traded)} \times \text{Average Price} $$

For example, suppose there are only three days in March – March 1st, March 2nd, and March 3rd.

If BYND (Beyond Meat) opens at \\$4 on March 1st and closes at \\$5 on March 3rd, its percent change for the month of March is $$\frac{\$5-\$4}{\$4} = +25.00\%$$

Suppose the high and low prices and volumes of BYND on each day are given below.
- March 1st: high \\$5, low \\$3, volume 500 million (0.5 billion)
- March 2nd: high \\$5.5, low \\$2.5, volume 1 billion
- March 3rd: high \\$5.25, low \\$4, volume 500 million (0.5 billion)

Then, the estimated total transaction volume is
$$\frac{\$5 + \$3}{2} \cdot 0.5 B + \frac{\$5.5 + \$2.5}{2} \cdot 1 B + \frac{\$5.25 + \$4}{2} \cdot 0.5 B = 8.3125B$$

```python
# base_url = f'https://financialmodelingprep.com/stable/historical-price-eod/full?symbol=AAPL'
# apikey = '?apikey=lXc6qZucwcvy741kF1HA6w6mRqv2inOo'
baseurl = 'https://financialmodelingprep.com/stable/historical-price-eod/full?symbol=AAPL&apikey=lXc6qZucwcvy741kF1HA6w6mRqv2inOo'
df = pd.DataFrame(requests.get(baseurl).json())
df['date'] = pd.to_datetime(df['date'])
filtered = df[(df['date'].dt.year == 2025) & (df['date'].dt.month == 6)].sort_values('date')
percent_change = (filtered.iloc[-1]['close'] - filtered.iloc[0]['open'])/(filtered.iloc[0]['open'])
if round(percent_change*100, 2) > 0:
    percent_change = '+' + str(round(percent_change*100, 2)) + '%'
else:
    percent_change = str(round(percent_change*100, 2)) + '%'
filtered['transaction_volume'] = filtered.apply(lambda df: (df['high'] + df['low'])/2 * df['volume'], axis=1)
total_transaction_volume = str(round(filtered['transaction_volume'].sum()/1_000_000_000, 2)) + 'B'
```

```python
# base_url = f'https://financialmodelingprep.com/stable/historical-price-eod/full?symbol=AAPL'
# apikey = '?apikey=lXc6qZucwcvy741kF1HA6w6mRqv2inOo'
base_url = 'https://financialmodelingprep.com/stable/historical-price-eod/full?symbol=AAPL&apikey=lXc6qZucwcvy741kF1HA6w6mRqv2inOo'
df = pd.DataFrame(requests.get(base_url).json())
df
```

Output:
```
symbol        date    open    high     low   close     volume  change  \
0      AAPL  2025-11-07  269.80  272.29  266.77  268.47   48279451   -1.32   
1      AAPL  2025-11-06  267.89  273.40  267.89  269.77   51204045    1.88   
2      AAPL  2025-11-05  268.61  271.70  266.93  270.14   43683100    1.53   
3      AAPL  2025-11-04  268.33  271.49  267.62  270.04   49274846    1.72   
4      AAPL  2025-11-03  270.42  270.85  266.25  269.05   50194600   -1.37   
...     ...         ...     ...     ...     ...     ...        ...     ...   
1251   AAPL  2020-11-13  119.44  119.67  117.87  119.26   81688586   -0.18   
1252   AAPL  2020-11-12  119.62  120.53  118.57  119.21  103350674   -0.41   
1253   AAPL  2020-11-11  117.19  119.63  116.44  119.49  112295000    2.30   
1254   AAPL  2020-11-10  115.55  117.59  114.13  115.97  138023400    0.42   
1255   AAPL  2020-11-09  120.50  121.99  116.05  116.32  154515315   -4.18   

      changePercent      vwap  
0          -0.49296  269.3325  
1   
```

```python
# # don't delete this cell, but do run it -- it is needed for the autograder tests

# # public test for stock_history
# history = stock_history('AAPL', 2025, 10)

# # public test for stock_stats
# stats = stock_stats(history)
```

```python
# grader.check("q3")
```

## Question 4 – Comment Threads 🧵

You regularly browse [Hacker News](https://news.ycombinator.com/) to keep up with the latest news in tech. One link to a Hacker News article is https://news.ycombinator.com/item?id=18344932. Note that this article has 18 comments and has a `storyid` of 18344932.

The problem now is that you don't have internet access on your phone during your morning commute to work, so you want to save the interesting stories' comment threads beforehand in a CSV. You find their [API documentation](https://github.com/HackerNews/API) and decide to get to work.

Complete the implementation of the function `get_comments`, which takes in a `storyid` and returns a DataFrame of all the comments below the news story. You can ignore "dead" comments(you will know them when you see them), as well as "dead" comments’ children. **Make sure the order of the comments in your DataFrame is from top to bottom just as you see on the website**. 

The DataFrame that `get_comments` returns should have 5 columns:
1. `'id'`: The unique ID of the comment.
2. `'by'`: The author of the comment.
3. `'text'`: The actual comment.
4. `'parent'`: The unique ID of the comment this comment is replying to.
5. `'time'`: When the comment was created (in `pd.Timestamp` format).

Some guidance:
1. The URL to make requests to is `'https://hacker-news.firebaseio.com/v0/item/{}.json'`, however, the `{}` should be replaced with the ID of the article or page you are trying to access. 
2. Again, do not `import json` – instead, use the `json` method on the Response object you get back.
3. Use depth-first search when traversing the comments tree. You will have to do this manually, since you cannot use Beautiful Soup (which is only for HTML documents, not JSON objects).
4. Make sure the length of your returned DataFrame is the same as value for the `'descendants'` key in the response JSON (both of which correspond to the number of comments for the story).
5. You are allowed to use loops in this function. You may also want to create at least one helper function.

<div class="alert alert-block alert-success">
    You may find <a href="https://www.youtube.com/watch?v=uOfwW-onmpc"><b>this hint video 🎥</b></a> helpful!
</div>

```python
story_id=18344932
url=f'https://hacker-news.firebaseio.com/v0/item/{story_id}.json'
requests.get(url).json()

rows = []

def get_item(item_id):
    url=f'https://hacker-news.firebaseio.com/v0/item/{item_id}.json'
    return requests.get(url).json()

def add_to_df(parsed_json):
    rows.append({
        'id': parsed_json.get('id'),
        'by': parsed_json.get('by'),
        'text': parsed_json.get('text'),
        'parent': parsed_json.get('parent'),
        'time': parsed_json.get('time')
    })
visited = set()

def dfs(item_id):
    if item_id in visited:
        return
    visited.add(item_id)

    parsed_json = get_item(item_id)
    if not parsed_json:
        return
    if parsed_json.get('dead'):
        return 
    
    if parsed_json.get('type') == 'story':
        for kid in parsed_json.get('kids', []):
            dfs(kid)
        return

    add_to_df(parsed_json)
    
    for kid in parsed_json.get('kids', []):
        dfs(kid)

dfs(story_id)

# rows = []
# stack = [story_id]
# visited = set()
# while stack:
#     curr = stack.pop()
#     if curr in visited:
#         continue
#     visited.add(curr)

#     parsed_json = get_item(curr)
#     if not parsed_json:
#         continue
#     if parsed_json.get('dead'):
#         continue
    
#     if parsed_json.get('type') == 'story':
#         stack.extend(parsed_json.get('kids', [])[::-1])
#     else:
#         add_to_df(parsed_json)
#         stack.extend(parsed_json.get('kids', [])[::-1])

df = pd.DataFrame(rows)
df['time'] = pd.to_datetime(df['time'], unit='s')
df
```

Output:
```
id           by                                               text  \
0   18380397      valyala  TimescaleDB is great for storing time series c...   
1   18346406       msiggy  I&#x27;m excited to give this database a try i...   
2   18348601      sman393  Can this be used side by side on normal Postgr...   
3   18348631   RobAtticus  Yep, absolutely. Regular PostgreSQL tables coe...   
4   18348984      sman393  Good to hear! how does the current TimescaleDB...   
5   18349540   RobAtticus  Not sure I follow exactly what you&#x27;re ask...   
6   18350673      sman393  Alright thanks! I thought I read that Timescal...   
7   18351061   RobAtticus  It does not support sharding writes across mul...   
8   18346750      zip1234  How fast is it when it has a TB of data? I rea...   
9   18347260      nevi-me  I spent about 8 months writing data to TSDB. I...   
10  18347555     dominotw  They have some numbers on their blog. Its very...   
11  18346476     dominotw  I evaluated this heavil
```

```python
# don't delete this cell, but do run it -- it is needed for the autograder tests
comments = get_comments(18344932)
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

## Congratulations! You're done Lab 6! 🏁

As a reminder, all of the work you want to submit needs to be in `lab.py`.

To ensure that all of the work you want to submit is in `lab.py`, we've included a script named `lab-validation.py` in the lab folder. You shouldn't edit it, but instead, you should call it from the command line (e.g. the Terminal) to test your work.

Once you've finished the lab, you should open the command line and run, in the directory for this lab:

```
python lab-validation.py
```

**This will run all of the `grader.check` cells that you see in this notebook, but only using the code in `lab.py` – that is, it doesn't look at any of the code in this notebook. If all of your `grader.check` cells pass in this notebook but not all of them pass in your command line with the above command, then you likely have code in your notebook that isn't in your `lab.py`!**

You can also use `lab-validation.py` to test individual questions. For instance,

```
python lab-validation.py q1 q2 q4
```

will run the `grader.check` cells for Questions 1, 2, and 4 – again, only using the code in `lab.py`. [This video](https://www.loom.com/share/0ea254b85b2745e59322b5e5a8692e91?sid=5acc92e6-0dfe-4555-9b6a-8115b6a52f99) how to use the script as well.

Once `python lab-validation.py` shows that you're passing all test cases, you're ready to submit your `lab.py` (and only your `lab.py`) to Gradescope. Once submitting to Gradescope, make sure to stick around until all test cases pass.

There is also a call to `grader.check_all()` below in _this_ notebook, but make sure to also follow the steps above.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q2 results: All test cases passed!

q3 results:
    q3 - 1 result:
        Trying:
            history.shape == (20, 13)
        Expecting:
            True
        **********************************************************************
        Line 1, in q3 0
        Failed example:
            history.shape == (20, 13)
        Exception raised:
            Traceback (most recent call last):
              File "c:\Users\sweek\miniforge3\envs\dsc80\Lib\doctest.py", line 1368, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q3 0[0]>", line 1, in <module>
                history.shape == (20, 13)
                ^^^^^^^
            NameError: name 'history' is not defined

    q3 - 2 result:
        Trying:
            history.label.iloc[-1] == 'June 03, 19'
        Expecting:
            True
        **********************************************************************
        Line 1, in q3 1
    
```

Tests for q1 (points: None)
```
>>> os.path.exists('lab06_1.html')
True
```
```
>>> q01_html = bs4.BeautifulSoup(open("lab06_1.html", encoding='utf-8'), features='lxml')
>>> len(q01_html.find_all('image')) == 0
True
```
failure_message: you should not use the <image> tag to add images


Tests for q2 (points: None)
```
>>> extract_book_out[1] == extract_book_url
True
```
```
>>> isinstance(get_product_info_out, dict)
True
```
```
>>> 'Category' in get_product_info_out.keys()
True
```
```
>>> get_product_info_out['Rating'] == 'Two'
True
```
```
>>> scrape_books_out.shape == (1, 11)
True
```
```
>>> scrape_books_out['Rating'][0] == 'Four'
True
```
```
>>> scrape_books_out['Title'][0] == 'Sharp Objects'
True
```


Tests for q3 (points: None)
```
>>> history.shape == (20, 13)
True
```
```
>>> history.label.iloc[-1] == 'June 03, 19'
True
```
```
>>> (len(stats[0]), len(stats[1])) == (7, 6)
True
```
```
>>> (float(stats[0][1:-1]) > 30) == True
True
```
```
>>> float(stats[1][:-1]) > 1 == True
True
```
```
>>> stats[1][-1] == 'B'
True
```


Tests for q4 (points: None)
```
>>> comments.shape == (18, 5)
True
```
```
>>> comments.loc[5, 'by'] == 'RobAtticus'
True
```
```
>>> comments.loc[5, 'time'].day == 31
True
```
