```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

# Lab 7 – Regular Expressions

## DSC 80, Fall 2025

### Due Date: Monday, November 17th at 11:59 PM

## Instructions

Welcome to the seventh DSC 80 lab this quarter!

Much like in DSC 10, this Jupyter Notebook contains the statements of the problems and provides code and Markdown cells to display your answers to the problems. Unlike DSC 10, the notebook is *only* for displaying a readable version of your final answers. The coding will be done in an accompanying `lab.py` file that is imported into the current notebook, and **you will only submit that `lab.py` file**, not this notebook!

Some additional guidelines:
- **Unlike in DSC 10, labs will have both public tests and hidden tests.** The bulk of your grade will come from your scores on hidden tests, which you will only see on Gradescope after the assignment deadline.
- **Do not change the function names in the `lab.py` file!** The functions in the `lab.py` file are how your assignment is graded, and they are graded by their name. If you changed something you weren't supposed to, you can find the original code in the [course GitHub repository](https://github.com/dsc-courses/dsc80-2025-sp).
- Notebooks are nice for testing and experimenting with different implementations before designing your function in your `lab.py` file. You can write code here, but make sure that all of your real work is in the `lab.py` file, since that's all you're submitting.
- You are encouraged to write your own additional helper functions to solve the lab, as long as they also end up in `lab.py`.

**To ensure that all of the work you want to submit is in `lab.py`, we've included a script named `lab-validation.py` in the lab folder. You shouldn't edit it, but instead, you should call it from the command line (e.g. the Terminal) to test your work.** More details on its usage are given at the bottom of this notebook.

**Importing code from `lab.py`**:

* Below, we import the `.py` file that's contained in the same directory as this notebook.
* We use the `autoreload` notebook extension to make changes to our `lab.py` file immediately available in our notebook. Without this extension, we would need to restart the notebook kernel to see any changes to `lab.py` in the notebook.
    - `autoreload` is necessary because, upon import, `lab.py` is compiled to bytecode (in the directory `__pycache__`). Subsequent imports of `lab` merely import the existing compiled python.

```python
%load_ext autoreload
%autoreload 2
```

Output:
```
The autoreload extension is already loaded. To reload it, use:
  %reload_ext autoreload
```

```python
from lab import *
```

```python
import pandas as pd
import numpy as np
import os
import re
```

## Question 1 – Practice with Regular Expressions 🛠
<div class="alert alert-block alert-warning">
<b>Note</b>: The functions in this question all have doctests in their docstrings. The doctests constitute the public tests; we will still run hidden tests on each of your functions. <b>Please don't change any of the docstrings!</b>
</div>

Regular expressions can be tricky, and the best way to gain familiarity with them is through lots of practice. In this question, you will work through ten exercises, each of which requires you to write a regular expression that matches strings that satisfy certain criteria. Make sure to take a close look at the doctests for each function in `lab.py`, as they provide useful guidance for the types of strings you should and shouldn't match.

***Notes:*** 
- Make sure to refer to the [Regular Expression Resources](https://dsc80.com/resources/#regular-expressions) posted on the course website. In particular, we recommend having [regex101.com](https://regex101.com/) open while working, along with the [cheat sheet](https://dsc80.com/resources/other/berkeley-regex-reference.pdf).

- Each exercise has a star rating, between 1 (⭐️) and 3 (⭐️⭐️⭐️) stars, indicating its difficulty level (1 being the easiest, 3 being the hardest). If you are spending lots of time on 1-star exercises, take a close look at the syntax from lecture, as there is probably an easier way of writing the necessary pattern!

<br>

### Exercise 1 (⭐️)

Write a regular expression that matches strings that have `'['` as the third character and `']'` as the sixth character.

<br>

### Exercise 2 (⭐️)

Write a regular expression that matches strings that are phone numbers that start with `'(858)'` and follow the format `'(xxx) xxx-xxxx'` (`'x'` represents a digit).

***Note:*** There is a space between `'(xxx)'` and `'xxx-xxxx'`.

<br>

### Exercise 3 (⭐️)

Write a regular expression that matches strings that:
- are between 6 and 10 characters long (inclusive),
- contain only alphanumeric characters, whitespace and `'?'`, and
- end with `'?'`.

<br>

### Exercise 4 (⭐️⭐️)

Write a regular expression that matches strings with exactly two `'$'`, one of which is at the start of the string, such that:
- the characters between the two `'$'` can be anything (including nothing) except the lowercase letters `'a'`, `'b'`, and `'c'`, (and `'$'`), and
- the characters after the second `'$'` can only be the **lowercase or uppercase** letters `'a'`/`'A'`, `'b'`/`'B'`, and `'c'`/`'C'`, with every `'a'`/`'A'` before every `'b'`/`'B'`, and every `'b'`/`'B'` before every `'c'`/`'C'`. There must be at least one `'a'` or `'A'`, at least one `'b'` or `'B'`, and at least one `'c'` or `'C'`.
    

<br>

### Exercise 5 (⭐️)
Write a regular expression that matches strings that represent valid Python file names, including the extension. 

***Note:*** For simplicity, assume that file names only contain letters, numbers, and underscores (`'_'`).

<br>

### Exercise 6 (⭐️)
Write a regular expression that matches strings that:
- are made up of only lowercase letters and exactly one underscore (`'_'`), and
- have at least one lowercase letter on both sides of the underscore.

<br>

### Exercise 7 (⭐️)
Write a regular expression that matches strings that start with and end with an underscore (`'_'`).

<br>

### Exercise 8 (⭐️)

Apple serial numbers are strings of length 1 or more that are made up of any characters, other than
- the uppercase letter `'O'`, 
- the lowercase letter `'i`', and 
- the number `'1'`.

Write a regular expression that matches strings that are valid Apple serial numbers.

<br>

### Exercise 9 (⭐️⭐️)

ID numbers are formatted as `'SC-NN-CCC-NNNN'`, where 
- SC represents state code in uppercase (e.g. `'CA'`),
- NN represents a number with 2 digits (e.g. `'98'`),
- CCC represents a three letter city code in uppercase (e.g. `'SAN'`), and
- NNNN represents a number with 4 digits (e.g. `'1024'`).

Write a regular expression that matches strings that are ID numbers corresponding to the cities of `'SAN'` or `'LAX'`, or the state of `'NY'`. Assume that there is only one city named `'SAN'`, only one city named `'LAX'`, and they are supposed to be in the `'CA'` state.

<br>

### Exercise 10 (⭐️⭐️⭐️)

Write a function named `match_10` that takes in a string and:
- converts the string to lowercase,
- removes all non-alphanumeric characters (i.e. removes everything that is not in the `\w` character class), and the letter `'a'`, and
- returns a list of every **non-overlapping** three-character substring in the remaining string, starting from the beginning of the string.
   
For instance, consider the following doctest:

```py
>>> match_10('Ab..DEF')
['bde']
```

Here's how `match_10` should process `'Ab..DEF'`:

1. Convert to lowercase: `'ab..def'`.
2. Remove non-alphanumeric characters and the letter `'a'`: `'bdef'`.
3. Starting from the beginning of the string, there is only a single non-overlapping three character substring: `'bde'`. Hence, we return `['bde']`.

***Note:*** Perform your operations in the exact order described above, otherwise your code may not pass all the tests. Don't use a `for`-loop.

```python

```

```python

```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

## Question 2 – Capturing Groups in Regular Expressions 📡

The dataset stored in `data/messy.txt` contains personal information from a fictional website that a user scraped from web server logs. Within this dataset, there are four fields that are of interest to you:
1. Email Addresses 
    - assume they are alphanumeric (dashes do not count) usernames and domain names
2. [Social Security Numbers](https://en.wikipedia.org/wiki/Social_Security_number#Structure)
3. Bitcoin Addresses (alphanumeric strings of long length)
4. Street Addresses

Complete the implementation of the function `extract_personal`, which takes in a string containing the contents of a server log file (like `open('data/messy.txt').read()`) and returns a **tuple of four separate lists** containing values of the 4 pieces of information listed above (in the order listed above). Do **not** keep empty values.

***Note:*** Since this data is messy, your function will be allowed to miss ~5% of the records in each list. Good spot checking using certain useful substrings (e.g. `'@'` for emails) should help assure correctness! Your function will be tested on a sample of the file `messy.txt`.

***Hint:*** There are multiple "delimiters" in use in the file; there are few enough of them that you can safely determine what they are.

```python
# experiment with extract_personal using the file s below
fp = os.path.join('data', 'messy.txt')
s = open(fp, encoding='utf8').read()
```

```python
email_pattern = r'[a-zA-Z0-9.]+@[a-zA-Z0-9.]+'
email_list = re.findall(email_pattern, s)
```

Output:
```
('lbeverstockc4@fda.gov',
 'zperkinl9@spotify.com',
 'bgergelydr@bing.com',
 'erowledge2s@wikia.com',
 'acarolinejz@angelfire.com',
 'mlangna@netlog.com',
 'sstroton83@blogger.com',
 'tmorritt6o@utexas.edu',
 'mnerval7f@tiny.cc',
 'twhetton77@msn.com',
 'khimpsonf2@instagram.com',
 'bwhitmelli6@reddit.com',
 'tgillis1n@skype.com',
 'cblavergc@umich.edu',
 'peslermc@reddit.com',
 'vgwynnecy@economist.com',
 'bpisco8q@aboutads.info',
 'jlemerchant5c@kickstarter.com',
 'sotsoncl@surveymonkey.com',
 'hsummonsly@rambler.ru',
 'ncorainijo@cam.ac.uk',
 'eoulettok@vk.com',
 'moxborrow5a@psu.edu',
 'rshearwoodrq@narod.ru',
 'dfernandesfa@cafepress.com',
 'retchinghamlh@prweb.com',
 'cchampa8@about.com',
 'drizzardodz@istockphoto.com',
 'dphibb52@netlog.com',
 'gpalffyle@jigsy.com',
 'mdibdall64@wikipedia.org',
 'sslatenis@slashdot.org',
 'kgoutcherqs@linkedin.com',
 'sknoles2j@sciencedirect.com',
 'aablittnu@economist.com',
 'ggeraldezr8@weather.com',
 'dottewell0@gnu.org',
 'ahuffeyc3@mac.com'
```

```python
social_security_pattern = r'\d{3}-\d{2}-\d{4}'
social_security_list = re.findall(social_security_pattern, s)
```

Output:
```
('721-18-4828',
 '891-94-5586',
 '736-45-4957',
 '493-15-4202',
 '198-82-6446',
 '589-07-4242',
 '864-68-1837',
 '741-97-0912',
 '353-94-6419',
 '194-23-9502',
 '742-87-8929',
 '263-56-3317',
 '170-10-1992',
 '872-79-0419',
 '315-72-8512',
 '129-15-4228',
 '790-15-3773',
 '767-39-9733',
 '820-59-9348',
 '741-31-7481',
 '366-01-9036',
 '691-36-6043',
 '421-03-8150',
 '107-15-0874',
 '794-78-0089',
 '395-10-2097',
 '319-83-2629',
 '826-73-1099',
 '684-98-6173',
 '378-21-1700',
 '850-03-5102',
 '814-90-4195',
 '335-62-3245',
 '375-79-4603',
 '250-40-1683',
 '496-12-1316',
 '716-67-0819',
 '277-93-7364',
 '880-40-5591',
 '668-02-0410',
 '724-03-8862',
 '709-10-1717',
 '780-05-0242',
 '831-10-4148',
 '724-18-1288',
 '896-21-0642',
 '844-91-3484',
 '368-48-5210',
 '442-64-0868',
 '105-14-4112',
 '158-84-1546',
 '609-81-5054',
 '845-06-0198',
 '549-38-9736',
 '391-57-0707',
 '233-02-5615',
 '342-81-1492',
 '137-05-2784',
 '437-40-9802',
 '100-61-7815',
 '791-39-5134',
 '603-08-9614',
 '417-05
```

```python
bitcoin_pattern = r'bitcoin:[a-zA-Z0-9]+'
bitcoin_list = re.findall(bitcoin_pattern, s)
bitcoin_list = [b.replace('bitcoin:', '') for b in bitcoin_list if b[-4:] != 'null']
```

Output:
```
('1Azt2u7uneLNUxukBP8YZicVmwTJit4hJ5',
 '153CgUdbUwBMUNHDcSDWHu1kyrZ9PJBNPL',
 '17EJQunKBnpADNq878Q7AxnTs43wzyJu9h',
 '1ED7inm1bSL6KXgZ8PmRRwWESe3k9nzhyr',
 '1KA4eqyUimncWJpCmBqFfCWz5GwKYLq5gd',
 '17z7gLz7xrhPw4CX51LXnh83XA8GaumDCu',
 '1AfC2ed4mtNbwpPsrncHShm1J7hAfB9PWw',
 '1FLuKwdTrCiMrdS4wmRKCJfTRDZsq5f2q2',
 '156UyXWH6qUausXjSzppVycMTFukkozXbh',
 '1AauxGneuh3c9WAm5dEtdy91pypVpW3CSq',
 '1AV61byrD55uxA4chVDE61V96fC3mKF7KT',
 '1C5JraJTyGiLUTVGf5T2zQVMLosPqteLix',
 '1EVCymahmT121vC7ZTxKs4eqWSfEjdbSfj',
 '1LRQTQKfRLgfa9QKb67qdmPSdx3RKqjdGm',
 '17EkvHmY17Y7dtJEgE82HynzEoB7ZaBD1k',
 '1xBbfoHGURrSyHSV74RRycLuwjhtUDAox',
 '18crm8zwe8KyN5Dfg5JDSVZCrvGikjKKto',
 '1GieJGbv2mbLKqmUiwnnpuiRMVdHviWe8a',
 '1MkW74jtuvCvC5U1N4koyq6HLKwwFMhUyR',
 '1GYTygPXmM29xWnQbpHyi8UoodWWAG45UE',
 '127jvkBRtEYGpkZE8K8wTfcijTwpVz5nYs',
 '187F7ZKKmitJnEyBSsGT3GSxga3hCwhuuN',
 '1DFZe1kyFmzivGGTJp4sWncGyzrqDzk3TD',
 '1CPUsbPvqa3ZF6NrFaQ1xjAtZBUbhDEdgK',
 '13zZF4NmoKDyD8mmTqPPwMfoR4H8cg2KC4',
 '18zShZdnezCqT936Qav1nMan
```

```python
street_address_pattern = r'\d+ \w+ \w+'
street_address_list = re.findall(street_address_pattern, s)
```

Output:
```
('66 Schurz Circle',
 '6953 Mesta Street',
 '918 Lunder Court',
 '2768 Waubesa Court',
 '24922 Gateway Pass',
 '27 Prentice Place',
 '886 Thackeray Road',
 '6303 Di Loreto',
 '6607 Sachs Way',
 '8 Hauk Crossing',
 '60387 Ryan Pass',
 '2216 Harbort Point',
 '3702 Ohio Center',
 '2 Delladonna Avenue',
 '320 Mayer Terrace',
 '67767 Merry Center',
 '276 1st Point',
 '57856 Mariners Cove',
 '0 Mccormick Center',
 '9842 Dapin Drive',
 '3620 Bartillon Parkway',
 '8485 Hanson Junction',
 '0303 Buhler Drive',
 '2412 Maryland Road',
 '4494 Golf Course',
 '786 Ludington Plaza',
 '238 Vera Court',
 '4 Delladonna Center',
 '564 Toban Pass',
 '7200 3rd Hill',
 '7381 Summit Center',
 '06 Loftsgordon Alley',
 '01 La Follette',
 '3 Monterey Alley',
 '8 Leroy Road',
 '4318 Service Circle',
 '83 North Way',
 '33 Carey Crossing',
 '55296 Gateway Avenue',
 '169 Claremont Point',
 '7 Starling Avenue',
 '9 Maple Junction',
 '82060 Spenser Trail',
 '8265 Jackson Avenue',
 '04653 Porter Alley',
 '364 Beilfuss 
```

```python
s.split('\n')
```

Output:
```
['1\t4/12/2018\tLorem ipsum dolor sit amet, consectetuer adipiscing elit. Proin risus. Praesent lectus.',
 '',
 'Vestibulum quam sapien| varius ut, blandit non, interdum in, ante. Vestibulum ante ipsum primis in faucibus orci luctus et ultrices posuere cubilia Curae; Duis faucibus accumsan odio. Curabitur convallis.|dottewell0@gnu.org\toR1mOq,!@#$%^&*(),[{bitcoin:18A8rBU3wvbLTSxMjqrPNc9mvonpA4XMiv\tIP:192.232.9.210\tccn:3563354617955160|ssn:380-09-9403}]|05-6609813,814 Monterey Court',
 '2\t12/18/2018\tSuspendisse potenti. In eleifend quam a odio. In hac habitasse platea dictumst.',
 '',
 'Maecenas ut massa quis augue luctus tincidunt. Nulla mollis molestie lorem. Quisque ut erat.,bassiter1@sphinn.com\tc5KvmarHX3o,test\u2060test\u202b,[{bitcoin:1EB7kYpnfJSqS7kUFpinsmPF3uiH9sfRf1,IP:20.73.13.197|ccn:3542723823957010\tssn:118-12-8276}#{bitcoin:1E5fev4boabWZmXvHGVkHcNJZ2tLnpM6Zv*IP:238.206.212.148\tccn:337941898369615,ssn:427-22-9352}#{bitcoin:1DqG3WcmGw74PjptjzcAmxGFuQdvWL7RCC,IP:171.241
```

```python
# don't change this cell, but do run it -- it is needed for the tests
test_fp = os.path.join('data', 'messy.test.txt')
test_s = open(test_fp, encoding='utf8').read()
emails, ssn, bitcoin, addresses = extract_personal(test_s)
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

## Question 3 – TF-IDF 📊

The dataset `data/reviews.txt` contains [Amazon reviews](https://cseweb.ucsd.edu/~jmcauley/datasets.html#amazon_reviews/) for ~200k phones and phone accessories. The dataset has already been "cleaned" for you. In this question, you will complete the implementation of a function, which takes in the reviews dataset as a Series (with one entry per review) as well as a single review, and returns the word that "best summarizes the single review" using TF-IDF.

To do so, implement the two functions below.

#### `tfidf_data`
You will first need to decompose the review into individual words, count their frequencies, and calculate the TF-IDF score for each word. Complete the implementation of the function `tfidf_data`, which takes in the reviews data as a Series (`reviews_ser`) and a single review (`review`) from that Series and returns a DataFrame indexed by the words in `review` with four columns:
- `'cnt'`: the number of times each word is found in the review (telling how frequently a word is used in the `review`)
- `'tf'`: the term frequency for each word (giving an idea of how important each word relative to the other words in the `review`)
- `'idf'`: the inverse document frequency for each word (helping understand how important each word is to the whole `reviews_ser`)
- `'tfidf'` the TF-IDF for each word (measuring how important each word is to `review`, taking into account both the frequency of the word in the `review` and across `reviews_ser`)

You may use a `for`-loop. The words in the outputted DataFrame may appear in any order.

You may also assume that the provided `review` will always be an element of the `reviews_ser` Series, so you do not need to account for division-by-zero errors.

***Hint:*** You may need to use the [`'\b'` character](https://www.regular-expressions.info/wordboundaries.html) somewhere.
    
<br>

#### `relevant_word`

Complete the implementation of the function `tfidf_data`, which takes in the DataFrame that `tfidf_data` returns and returns the word that "best summarizes" the review. If there are multiple "best" summary words, return any one of them.



```python
# experiment with tfidf_data using reviews_ser and review below 
fp = os.path.join('data', 'reviews.txt')
reviews_ser = pd.read_csv(fp, header=None).squeeze("columns")
review = open(os.path.join('data', 'review.txt'), encoding='utf8').read().strip()
```

```python
review = review.lower()
split_words = review.split()
unique_words = np.unique(split_words)
review_data = pd.DataFrame(columns=['cnt', 'tf', 'idf', 'tfidf'])

for word in unique_words:
    re_word = fr'\b{word.lower()}\b'
    cnt = len(re.findall(re_word, review))
    tf = cnt / len(split_words)
    idf = np.log(len(reviews_ser) / (reviews_ser.str.contains(re_word).sum()))
    tfidf = tf * idf
    review_data.loc[word] = [cnt, tf, idf, tfidf]

review_data

```

Output:
```
cnt        tf       idf     tfidf
5                1.0  0.011765  2.271546  0.026724
a                4.0  0.047059  0.360392  0.016960
against          1.0  0.011765  4.266529  0.050194
also             2.0  0.023529  1.924159  0.045274
and              5.0  0.058824  0.248188  0.014599
before           1.0  0.011765  2.820618  0.033184
case             3.0  0.035294  1.121935  0.039598
chunk            1.0  0.011765  8.028953  0.094458
color            1.0  0.011765  2.835119  0.033354
combinations     1.0  0.011765  7.259820  0.085410
comes            1.0  0.011765  2.829584  0.033289
cover            2.0  0.023529  2.682848  0.063126
covers           1.0  0.011765  3.747910  0.044093
create           1.0  0.011765  5.814336  0.068404
damage           1.0  0.011765  4.285037  0.050412
design           1.0  0.011765  3.055127  0.035943
different        2.0  0.023529  2.960048  0.069648
from             1.0  0.011765  1.605962  0.018894
great            1.0  0.011765  1.301647  0.0153
```

```python
review_data['tfidf'].idxmax()
```

Output:
```
np.str_('chunk')
```

```python

```

```python
# don't change this cell, but do run it -- it is needed for the tests
fp = os.path.join('data', 'reviews.txt')
reviews_ser = pd.read_csv(fp, header=None).squeeze("columns")
review = open(os.path.join('data', 'review.txt'), encoding='utf8').read().strip()
q3_tfidf = tfidf_data(reviews_ser, review)

try:
    q3_rel = relevant_word(q3_tfidf)
except:
    q3_rel = None
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

## Questions 4 and 5 – Tweet Analysis 🐥

The dataset `data/ira.csv` contains tweets tagged by Twitter as likely being posted by the [Internet Research Agency](https://en.wikipedia.org/wiki/Internet_Research_Agency), the tweet factory facing allegations for attempting to influence US political elections.

Questions 4 and 5 will focus on the following:
- Question 4: Look at the hashtags present in the text and trends in their makeup.
- Question 5: Prepare the dataset for modeling by creating features out of the text fields.

### Question 4 – Hashtags #️⃣

You may assume that a hashtag is any string without whitespace that immediately follows a `'#'` symbol. To clarify, a hashtags can contain a `'#'` symbol, as long as there is no whitespace. For example, `'#data#science'` is considered one individual hashtag, whereas `'#data #science'` is considered two individual hashtags.

#### `hashtag_list`

Complete the implementation of the function `hashtag_list`, which takes in a Series of tweet texts and returns a Series containing a list of hashtags present in each tweet's text. If a tweet's text doesn't contain a hashtag, the Series should contain an empty list for that tweet. Don't include the leading `'#'` symbol in the lists that are returned.

<br>

#### `most_common_hashtag`

Complete the implementation of the function `most_common_hashtag`, which takes in a Series of hashtag lists (as is outputted by `hashtag_list`) and returns a Series consisting of a single hashtag per tweet: 
- If the tweet's text has no hashtags, the entry should in the output Series should be `NaN`.
- If the tweet's text has one distinct hashtag, the entry in the output Series should be that hashtag.
- If the tweet's text has more than one hashtag, the entry in the output Series should be the most common hashtag **in the tweet's text** with respect to **the whole input Series**. If there is a tie for the most common, any of the most common can be returned.
    - For example, if the input Series was `pd.Series([[2], [2], [3, 2, 3]])`, the output would be `pd.Series([2, 2, 2])`. Even though `3` was more common in the third list than `2`, `2` is more common than `3` among all hashtags in the Series.

```python
# The public tests don't test your work on the `ira` data,
# but the hidden tests do.
# So, make sure to thoroughly test your work yourself!
fp = os.path.join('data', 'ira.csv')
ira = pd.read_csv(fp, names=['id', 'name', 'date', 'text'])
```

```python
hashtag_pattern = r'#\S+'
hashtags = ira['text'].apply(lambda x: re.findall(hashtag_pattern, x))
hashtags = hashtags.apply(lambda lst: [tag[1:] for tag in lst])
hashtag_counts = {}

for hashtag in hashtags.explode():
    hashtag_counts[hashtag] = hashtag_counts.get(hashtag, 0) + 1

def get_most_common(lst):
    if not lst:
        return np.nan
    if len(lst) == 1:
        return lst[0]
    most_common = max(lst, key=lambda x: hashtag_counts[x])
    return most_common

most_common = hashtags.apply(get_most_common)
most_common
```

Output:
```
0                        CatTV
1                          NaN
2                         tech
3                         news
4        IHatePokemonGoBecause
                 ...          
89995                 politics
89996      ThingsYouCantIgnore
89997                      NaN
89998             rantfortoday
89999                     news
Name: text, Length: 90000, dtype: object
```

```python

```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

### Question 5 – Features 📋

Now, create a DataFrame of features from the `ira` data.  That is, create a function `create_features` that takes in a DataFrame `ira` that has just a single column, `'text'`, and returns a DataFrame with the same index as `ira` (i.e. the rows correspond to the same tweets) and the following columns:
* `'num_hashtags'`, the number of hashtags present in the tweet.
* `'mc_hashtags'`, the most common hashtag associated to the tweet (using the result of `most_common_hashtag` from Question 4).
* `'num_tags'`, the number of tags the tweet has.
    - A tag is a string starting with the `'@'` symbol followed by only alphanumerical characters. A`'@'` symbol on its own does not count as a tag.
* `'num_links'`, the number of hyperlinks present in the tweet.
    - A hyperlink is a string starting with `'http://'` or `'https://'`, not followed by whitespaces.
* `'is_retweet'`, a Boolean describing whether the tweet is a retweet. A retweet is a tweet that **begins** with `'RT'`.
* `'text'`, a version of the tweet's text that is cleaned according to the following steps, **in this exact order**:
    1. All meta-information above (retweet info, tags, hyperlinks, and hashtags) should be replaced with a single space.
    2. Everything other than letters, numbers, and spaces should be replaced with a single space.
    3. All letters should be lowercase.
    4. All words should be separated by exactly one space, and leading/trailing whitespace should be removed (stripped).
    
The columns in the outputted DataFrame must be in the order `['text', 'num_hashtags', 'mc_hashtags', 'num_tags', 'num_links', 'is_retweet']`. (Remember, the DataFrame that `create_features` is called on only has a single column, `'text'`.)

***Notes:***
- It's a good idea to make helper function for each column.
- The `\w` character class in regex **does not** refer to letters, numbers, and spaces (or even just letters and numbers). As such, you can't use it here!
- `create_features` will take a while to run on the entire dataset – test it on a small sample first!
- A `'@'` by itself is not a tag - do not account for empty strings as tags

```python
# The doctests/public tests don't test your work on the `ira` data,
# but the hidden tests do.
# So, make sure to thoroughly test your work yourself!
fp = os.path.join('data', 'ira.csv')
ira = pd.read_csv(fp, names=['id', 'name', 'date', 'text'])
```

```python
hashtag_pattern = r'#\S+'
tag_pattern = r'@\S+'
hyperlink_pattern = r'https?://\S+'
num_hashtags = ira['text'].apply(lambda x: len(re.findall(hashtag_pattern, x)))
mc_hashtags = most_common_hashtag(hashtag_list(ira['text']))
num_tags = ira['text'].apply(lambda x: len(re.findall(tag_pattern, x)))
num_links = ira['text'].apply(lambda x: len(re.findall(hyperlink_pattern, x)))
is_retweet = ira['text'].str.startswith('RT')
```

Output:
```
0        False
1         True
2        False
3        False
4         True
         ...  
89995    False
89996     True
89997     True
89998     True
89999    False
Name: text, Length: 90000, dtype: bool
```

```python
def cleanup_text(text):
    hashtag_pattern = r'#\S+'
    tag_pattern = r'@\S+'
    hyperlink_pattern = r'https?://\S+'
    non_text_pattern = r'[^a-zA-Z0-9 ]'
    # Step 1: Remove retweet information, hashtags, tags, and hyperlinks
    text = text.replace('RT ', '')
    text = re.sub(hashtag_pattern, '', text)
    text = re.sub(tag_pattern, '', text)
    text = re.sub(hyperlink_pattern, '', text)
    # Step 2: Remove non-text characters
    text = re.sub(non_text_pattern, '', text)
    # Step 3: Convert to lowercase
    text = text.lower()
    # Step 4: Remove extra spaces
    text = ' '.join(text.split())
    text = text.strip()
    return text

text = ira['text'].apply(cleanup_text)
text
```

Output:
```
0           the best exercise to lose belly fat in 2 weeks
1        dozens of hate groups have charity status chro...
2        artificial intelligence can find map poverty r...
3        uber balks at rules proposed by worlds busiest...
4        he didnt let me do that for a klondike bar scr...
                               ...                        
89995              trump kasich shouldnt be allowed to run
89996               the last step at the top of the stairs
89997    when someone said the first link from in my ra...
89998    i speak the word of god therefore because the ...
89999                         10 things to know for monday
Name: text, Length: 90000, dtype: object
```

```python
# don't change this cell, but do run it -- it is needed for the tests
# (yes, we know it says "hidden" – there are still truly hidden tests in this question)
fp_hidden = 'data/ira_test.csv'
ira_hidden = pd.read_csv(fp_hidden, header=None)
text_hidden = ira_hidden.iloc[:, -1:]
text_hidden.columns = ['text']

test_hidden = create_features(text_hidden)
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

## Congratulations! You're done Lab 7! 🏁

As a reminder, all of the work you want to submit needs to be in `lab.py`.

To ensure that all of the work you want to submit is in `lab.py`, we've included a script named `lab-validation.py` in the lab folder. You shouldn't edit it, but instead, you should call it from the command line (e.g. the Terminal) to test your work.

Once you've finished the lab, you should open the command line and run, in the directory for this lab:

```
python lab-validation.py
```

**This will run all of the `grader.check` cells that you see in this notebook, but only using the code in `lab.py` – that is, it doesn't look at any of the code in this notebook. If all of your `grader.check` cells pass in this notebook but not all of them pass in your command line with the above command, then you likely have code in your notebook that isn't in your `lab.py`!**

You can also use `lab-validation.py` to test individual questions. For instance,

```
python lab-validation.py q1 q2 q4
```

will run the `grader.check` cells for Questions 1, 2, and 4 – again, only using the code in `lab.py`. [This video](https://www.loom.com/share/0ea254b85b2745e59322b5e5a8692e91?sid=5acc92e6-0dfe-4555-9b6a-8115b6a52f99) how to use the script as well.

Once `python lab-validation.py` shows that you're passing all test cases, you're ready to submit your `lab.py` (and only your `lab.py`) to Gradescope. Once submitting to Gradescope, make sure to stick around until all test cases pass.

There is also a call to `grader.check_all()` below in _this_ notebook, but make sure to also follow the steps above.

```python

```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q2 results: All test cases passed!

q3 results: All test cases passed!

q4 results: All test cases passed!

q5 results: All test cases passed!
```

Tests for q1 (points: None)
```
>>> import doctest
>>> doctest.run_docstring_examples(match_1, {'match_1': match_1})
>>> doctest.run_docstring_examples(match_2, {'match_2': match_2})
>>> doctest.run_docstring_examples(match_3, {'match_3': match_3})
>>> doctest.run_docstring_examples(match_4, {'match_4': match_4})
>>> doctest.run_docstring_examples(match_5, {'match_5': match_5})
>>> doctest.run_docstring_examples(match_6, {'match_6': match_6})
>>> doctest.run_docstring_examples(match_7, {'match_7': match_7})
>>> doctest.run_docstring_examples(match_8, {'match_8': match_8})
>>> doctest.run_docstring_examples(match_9, {'match_9': match_9})
>>> doctest.run_docstring_examples(match_10, {'match_10': match_10})
```
failure_message: doctests


Tests for q2 (points: None)
```
>>> emails[0] == 'test@test.com'
True
```
```
>>> ssn[0] == '423-01-9575'
True
```
```
>>> bitcoin[0] == '1BvBMSEYstWetqTFn5Au4m4GFg7xJaNVN2'
True
```
```
>>> addresses[0] == '530 High Street'
True
```


Tests for q3 (points: None)
```
>>> bool(q3_tfidf['cnt'].sum() == 85)
True
```
```
>>> 'before' in q3_tfidf.index
True
```
```
>>> bool(q3_tfidf.loc['a', 'cnt'] == 4)
True
```
```
>>> bool(np.isclose(q3_tfidf.loc['phone', 'idf'], 0.820703, atol=0.5))
True
```
```
>>> q3_rel in q3_tfidf.index
True
```


Tests for q4 (points: None)
```
>>> testdata = [['RT @DSC80: Text-cleaning is cool! #NLP https://t.co/xsfdw88d #NLP1 #NLP1']]
>>> test = pd.DataFrame(testdata, columns=['text'])['text']
>>> out = hashtag_list(test)
>>> out.iloc[0] == ['NLP', 'NLP1', 'NLP1']
True
```
```
>>> testdata = [['RT @DSC80: Text-cleaning is cool! #NLP https://t.co/xsfdw88d #NLP1 #NLP1']]
>>> test = hashtag_list(pd.DataFrame(testdata, columns=['text'])['text'])
>>> most_common_hashtag(test).iloc[0] == 'NLP1'
True
```
```
>>> out = pd.Series([[2], [2], [3, 2, 3]])
>>> bool((most_common_hashtag(out) == 2).all())
True
```


Tests for q5 (points: None)
```
>>> testdata = [['RT @DSC80: Text-cleaning is cool! #NLP https://t.co/xsfdw88d #NLP1 #NLP1']]
>>> test = pd.DataFrame(testdata, columns=['text'])
>>> out = create_features(test)
>>> anscols = ['text', 'num_hashtags', 'mc_hashtags', 'num_tags', 'num_links', 'is_retweet']
>>> ansdata = [['text cleaning is cool', 3, 'NLP1', 1, 1, True]]
>>> ans = pd.DataFrame(ansdata, columns=anscols)
>>> bool((out == ans).all().all())
True
```
```
>>> test_hidden.columns.tolist() == ['text', 'num_hashtags', 'mc_hashtags', 'num_tags', 'num_links', 'is_retweet']
True
```
```
>>> bool(400 <= (test_hidden['num_tags'] == 0).sum() <= 450)
True
```
```
>>> bool(450 <= (test_hidden['num_links'] == 0).sum() <= 550)
True
```
